Before you turn this problem in, make sure everything is saved and runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE`, as well as your name below:

```python
NAME = "Mehmet ÖRS"
```

Remove the lines with "raise NotImplementedError()" after you implemented the given part of the code.

Do **not** remove the lines with "# YOUR CODE HERE".

---

# Machine Learning Bootcamp - Coding Challenge

Thank you for applying to Google Developers Machine Learning Bootcamp! We are thrilled to have you as a potential candidate for our program.

The application process includes two sections designed to assess your skills and knowledge in machine learning. The first section will test your ability to write algorithms, while the second section will evaluate your understanding of fundamental concepts related to machine learning.

Our bootcamp aims to provide participants with a comprehensive understanding of machine learning and equip them with the practical skills needed to succeed in this field. The program is highly selective, and we are seeking candidates who are passionate about machine learning and have a strong foundation in mathematics, statistics, and computer science.

In the first section, you will be required to write code for a single algorithmic problem to assess your problem-solving and coding abilities. In the second section, we will assess your comprehension of fundamental mathematical concepts related to machine learning and your proficiency in translating them into code.

We understand that the application process can be challenging, but we encourage you to showcase your strengths and potential as a machine learning practitioner. Our team is committed to selecting candidates who demonstrate intellectual curiosity, a strong work ethic, and a collaborative mindset.

We look forward to reviewing your application.

<img src="https://windows10spotlight.com/wp-content/uploads/2019/04/2ddf44eb7f24fea177d6ac024129442a.jpg">

Reference: Panoramic flower gardens of Shikisai-no-oka hill, Biei, Furano basin, Hokkaido, Japan

### Problem Statement
Given a garden with different types of flowers, try to find if a particular combination of flowers is present. Each flower in the garden is represented by a character (D = Daisies, F = Daffodils, I = Irises, L = Lilies, M = Marigolds, P = Peonies, R = Roses, S = Sunflowers, T = Tulips, V = Violets). For example, consider the following garden:

<pre>
FILMPRSTVD
DVTSRPMLIF
FFFFFFFFFF
FFFFFFFFFF
IIIIIIIIII
</pre>

The combination of flowers to search for is:<pre>
TSRPML
FFFFFF
FFFFFF
</pre>

The combination begins at the second row and the third column of the grid and continues in the following two rows. The combination is said to be present in the garden. The return value should be 'YES' or 'NO', depending on whether the combination is found. In this case, return YES.

<br>

---

**Function Description** <br>
Complete the flower_search function in the editor below. It should return YES if the combination exists in the garden, or NO otherwise.

flower_search has the following parameter(s):

string G: the garden to search <br>
string P: the combination to search for

<br>

**Input Format** <br>
The test cases is represented as follows:
There are two input strings. 
The first string is the lines that represent the garden.
The second string is the lines that represent the combination.

<br> **Returns** <br>
string: either YES or NO

<br>

---

**Sample Input 1**

<pre>
SITLMPPTRM
RSLFFPTRFV
TVTTIMIRML
LTLDPTVLIM
IIIVPDPTFL
PRLLTMPLSM
RMSLPLDIVL
SDPLFDRRDF
DTLMITIVPR
MRDSVIMFLS

VPDP
LTMP
LPLD
</pre>

**Sample Output 1**
<pre>
YES
</pre>

**Explanation 1**

The combination is present in the larger garden as marked in bold below.
<pre>
SITLMPPTRM
RSLFFPTRFV
TVTTIMIRML
LTLDPTVLIM
III<b>VPDP</b>TFL
PRL<b>LTMP</b>LSM
RMS<b>LPLD</b>IVL
SDPLFDRRDF
DTLMITIVPR
MRDSVIMFLS
</pre>

**Sample Input 2**
<pre>
MDDMPLPVIFIRPRD
FFMIFLFLLDVTRVI
MSMLTRDTITSVRMT
PIILPRVPFFTVFRV
TTSFDVMPDMTSMVR
IPITDIRLLLTTSTI
PDISSFMTMVRRSMT
DSPVSPIDSRVLSTD
PFFSVVSTVPRITDR
MDMDDSMPMISIPDM
PMVDMLTDVVFRDTD
VRIMFDTDVPLMTFF
MMPTVLPILSLLMSP
SRTSDPLDLIFMFSM
RPDRIVISDTTSFRD

VV
VV
</pre>

**Sample Output 2**
<pre>
NO
</pre>

**Explanation 2**

The combination is not found in the larger garden.

### Solution

<!-- grade_id: cell-6500f8208bc77e4b, points: 0, grade: false, solution: true -->
```python
# Do not import any additional libraries
import math
import os
import random
import re
import sys

# Complete the 'flower_search' function below.
#
# The function is expected to return a STRING.
# The function accepts following parameters:
#  1. STRING G
#  2. STRING P

def flower_search(G, P):
     # YOUR CODE HERE
  
  first_chars = {s[0] for s in P.split("\n")}
  flower_sequences = P.split("\n")

  G = G.replace("\n", '')
  
  for i in range(len(G)):
    if G[i] in first_chars:
      #search which sequence starts with that flower
      for seq in flower_sequences:
        #if the sequence being checked is longer than the rest of the garden, return no
        if len(seq) > len(G) - i:
          return "NO"
        if seq[0] == G[i]:
          j = i
          found = True
          for flower in seq:
            if G[j] != flower:
              found = False
              break
            j += 1

          #pop the flower if it is found
          if found:
            flower_sequences.remove(seq)

          #continue fro the end of the found flower
  

  if len(flower_sequences) == 0:
    return "YES"

  return "NO"
   
```

<!-- grade_id: cell-609405b2f549b5b6, points: 50, grade: true, solution: false -->
```python
# DO NOT CHANGE THIS CELL

test_garden = "SITLMPPTRM\nRSLFFPTRFV\nTVTTIMIRML\nLTLDPTVLIM\nIIIVPDPTFL\nPRLLTMPLSM\nRMSLPLDIVL\nSDPLFDRRDF\nDTLMITIVPR\nMRDSVIMFLS"
test_flower_1 = "VPDP\nLTMP\nLPLD"
test_flower_2 = "VV\nVV"

assert flower_search(test_garden, test_flower_1) == "YES"
assert flower_search(test_garden, test_flower_2) == "NO"
```

## Part 2: Fundamental ML Concept Understanding
Reference: https://machinelearningmastery.com/implement-simple-linear-regression-scratch-python/

<img src="https://imgs.xkcd.com/comics/linear_regression.png">

Reference: https://xkcd.com/1725/

### Problem Statement

You have been tasked with implementing a basic linear regression model using only Python core libraries. Your task is to write a function that takes in two lists, X and y, where X represents the features and y represents the target variable. Your function should return the slope (m) and intercept (b) that minimize the sum of squared errors (SSE) between the predicted and actual target values.

To implement the model, you can use the following steps:

- Calculate the mean of X and y using a for loop
- Calculate the variance of X using a for loop
- Calculate the covariance between X and y using a for loop
- Calculate the slope m using the formula: m = cov(X, y) / var(X)
- Calculate the intercept b using the formula: b = mean(y) - m * mean(X)

To test the function, you can create sample data as follows:

<pre>
X = [1, 2, 3, 4, 5]
y = [2.1, 3.9, 6.2, 8.1, 9.8]

m, b = linear_regression(X, y)

print(f"Slope (m): {m:.2f}")
print(f"Intercept (b): {b:.2f}")
</pre>

The output should be:

<pre>
Slope (m): 1.98
Intercept (b): 0.12
</pre>

These values represent the slope and intercept of the line that best fits the sample data, minimizing the sum of squared errors.

### Solution

<!-- grade_id: cell-be4823434bcf4d18, points: 0, grade: false, solution: true -->
```python
def mean(values):
    """
    Calculate the mean of a list of values.

    Parameters:
    values (list): A list of numeric values.

    Returns:
    float: The mean value.
    """
    # YOUR CODE HERE

    sum = 0
    for value in values:
      sum += value
    return sum/len(values) 
    
  
```

<!-- grade_id: cell-ab48ac1f5f8a01f2, points: 10, grade: true, solution: false -->
```python
# DO NOT CHANGE THIS CELL

sample_X = [1, 2, 3, 4, 5]
sample_y = [2.1, 3.9, 6.2, 8.1, 9.8]

assert round(mean(sample_X), 2) == 3.0
assert round(mean(sample_y), 2) == 6.02

```

<!-- grade_id: cell-08a4f7604dac88d8, points: 0, grade: false, solution: true -->
```python
def variance(values):
    """
    Calculate the variance of a list of values.

    Parameters:
    values (list): A list of numeric values.

    Returns:
    float: The variance value.
    """
    # YOUR CODE HERE
   
    
    mean_val = mean(values)
    sum = 0
    for xi in values:
      sum += (xi - mean_val)**2 

    return sum / (len(values))
    
    raise NotImplementedError()

```

<!-- grade_id: cell-1c89306425776515, points: 10, grade: true, solution: false -->
```python
# DO NOT CHANGE THIS CELL

sample_X = [1, 2, 3, 4, 5]
sample_y = [2.1, 3.9, 6.2, 8.1, 9.8]

assert round(variance(sample_X), 2) == 2
assert round(variance(sample_y), 2) == 7.7

```

<!-- grade_id: cell-7a5d8312d93e6392, points: 0, grade: false, solution: true -->
```python
def covariance(x_values, y_values):
    """
    Calculate the covariance between two lists of values.

    Parameters:
    x_values (list): A list of numeric values.
    y_values (list): A list of numeric values.

    Returns:
    float: The covariance value.
    """
    # YOUR CODE HERE

    xMean = mean(x_values)
    yMean = mean(y_values)

    sum = 0
    for i in range(len(x_values)):
      sum += (x_values[i] - xMean) * (y_values[i] - yMean)


    return sum / len(x_values)

```

<!-- grade_id: cell-1fb53a2266efc731, points: 10, grade: true, solution: false -->
```python
# DO NOT CHANGE THIS CELL

sample_X = [1, 2, 3, 4, 5]
sample_y = [2.1, 3.9, 6.2, 8.1, 9.8]

assert round(covariance(sample_X, sample_y), 2) == 3.92

```

<!-- grade_id: cell-603d1f44da344323, points: 0, grade: false, solution: true -->
```python
def linear_regression(X, y):
    """
    Implement a basic linear regression model using Python core libraries.

    Parameters:
    X (list): A list of features.
    y (list): A list of target values.

    Returns:
    float: The slope (m) that minimizes the sum of squared errors.
    float: The intercept (b) that minimizes the sum of squared errors.
    """
    # YOUR CODE HERE
    
    m = covariance(X, y) / variance(X)
    b = mean(y) - m * mean(X)

    
    return m, b

```

<!-- grade_id: cell-9ca37d3365e1d0da, points: 20, grade: true, solution: false -->
```python
# DO NOT CHANGE THIS CELL

sample_X = [1, 2, 3, 4, 5]
sample_y = [2.1, 3.9, 6.2, 8.1, 9.8]

assert round(linear_regression(sample_X, sample_y)[0], 2) == 1.96
assert round(linear_regression(sample_X, sample_y)[1], 2) == 0.14

```